<a href="https://colab.research.google.com/github/mmanam2/mmanam2/blob/main/Week03/week03_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 3 Exercises

See: _McKinney 2.3_ and [Python Documentation](https://docs.python.org/3/tutorial/controlflow.html) section 4 on flow control.


**At the begining of the semester, all of the workshop programming exercises will be structured a specific way to make it easier to verify for yourself that you're on the right track as well as easier for me to do a first pass on automated grading.  The structure of each question will require you to write a function using Python code. Don't worry that we haven't talked about functions yet. Just edit the code between** `### BEGIN SOLUTION` and `### END SOLUTION` **as shown in the example below.**


**WHAT I PROVIDE:**
```
def some_function(parameter1, parameter2):

   ### BEGIN SOLUTION
   x = -1
   ### END SOLUTION
   
   return x
```

**WHAT YOU SHOULD DO:** Just change the parameter names (if you feel you need to) and the calculations between `### BEGIN SOLUTION` and `### END SOLUTION`.  This is just a made up example.
```
def some_function(a, b):

   ### BEGIN SOLUTION
   temp = a + b
   x = temp / a * b
   ### END SOLUTION
   
   return x
```


---
---

**Below each programming exercise are some tests (`assertions`) that verify your code is working correctly.  If any assertions fail, you know that something isnt' right with your code, but having all assertions pass doesn't necessarily mean your code is perfect, yet. You should also create your own tests to make sure your code is correct.**

**For now, please don't change any function names**

### 14.1 Difference in rate per 1,000

Often in public health, we report metrics as a number per 1,000 or per 1,000,000 in population. The purpose of that is to normalize the numbers between area of larger and smaller populations.

Below, we have a function already built to calculate teh different in rate per 1,000 in two different regions.  The function normalizes them and then prints out a message describing their difference.

```python
def diff_in_rate_normalized(count_a, total_a, count_b, total_b):
    """(int,int,int,int) -> str
    * count_a is the number of occurences in area A
    * total_a is the total population in area A
    * count_b is the number of occurences in area B
    * total_b is the total population in area B

    This function returns a string describing how A and B compare in terms of occurence rates per 1,000 population."""

    rate_a = count_a / total_a * 1000
    rate_b = count_b / total_b * 1000

    if rate_a == rate_b:
        msg = "The rate in A and the rate in B are the same ({}).".format(rate_a)
    elif rate_a > rate_b:
        msg = "The rate in A ({}) is greater than the rate in B ({}).".format(rate_a, rate_b)
    else:
        msg = "The rate in A ({}) is less than the rate in B ({})".format(rate_a, rate_b)

    return msg
```

```python
rate_covid_testing = diff_in_rate_normalized(1,1000,15,10000)
print(rate_covid_testing)
```

Output:
```
The rate in A (1.0) is less than the rate in B (1.5)
```

```python
diff_in_rate_normalized(3,1000,30,10000)
```

Output:
```
'The rate in A and the rate in B are the same (3.0).'
```

### 14.2 Trimming outliers

While it works well in most cases, in practice, this may not truly make numbers comparable between extremely large population centers (e.g. New York City at 8.5 million) and very small rural areas (e.g. Meeteese, WY at 459 people).

Let's take that function and make some adjustments.  If the total population of either A or B is more than 100 times larger than the other, then we want to return a different message.  That is, if the populations are more than two orders of magnitude different, then we shouldn't try to compare them.

```python
def diff_in_rate_normalized(count_a, total_a, count_b, total_b):
    """
    (int, int, int, int) -> str
    * count_a is the number of occurrences in area A
    * total_a is the total population in area A
    * count_b is the number of occurrences in area B
    * total_b is the total population in area B

    This function returns a string describing how A and B compare in terms of occurrence rates per 1,000 population.

    If total_a / total_b > 100 OR total_b / total_a > 100 then we'll return a message saying the two can't be compared.
    "The total populations in A and B are so different that they can't be compared."
    """

    # Check for significant population difference
    if total_a / total_b > 100 or total_b / total_a > 100:
        return "The total populations in A and B are so different that they can't be compared."

    # Calculate the occurrence rates per 1,000 population
    rate_a = (count_a / total_a) * 1000
    rate_b = (count_b / total_b) * 1000

    # Create a comparison message
    if rate_a > rate_b:
        msg = f"Area A has a higher occurrence rate ({rate_a:.2f} per 1,000) than Area B ({rate_b:.2f} per 1,000)."
    elif rate_b > rate_a:
        msg = f"Area B has a higher occurrence rate ({rate_b:.2f} per 1,000) than Area A ({rate_a:.2f} per 1,000)."
    else:
        msg = "Area A and Area B have the same occurrence rate."

    return msg


```

```python

# Now, calling the function with the provided test case
result = diff_in_rate_normalized(5, 459, 30, 8500)
print(result)
```

Output:
```
Area A has a higher occurrence rate (10.89 per 1,000) than Area B (3.53 per 1,000).
```

### 14.3 LACE Score
There is a simple readmission index called the LACE Score: https://www.hindawi.com/journals/bmri/2015/169870/tab1/

Use this documentation to create a function that can compute LACE scores based on the 4 input parameters.  Below is the function signature and documentation that you start with.

<!-- grade_id: bmi-answer, points: 0, grade: false, solution: true -->
```python
def LACE(length_of_stay, acute_flag, charlson, ed_visits):
    """(int, bool, int, int) -> int
    This function uses the logic from https://www.hindawi.com/journals/bmri/2015/169870/tab1/
    to compute the LACE score for this patient.

    >>> LACE(4, False, 1, 0)
    5

    >>> LACE(4, True, 4, 7)
    16

    """

    # Length of Stay (LOS) points
    if length_of_stay <= 3:
        los_score = length_of_stay
    elif 4 <= length_of_stay <= 6:
        los_score = 4
    elif 7 <= length_of_stay <= 13:
        los_score = 5
    else:
        los_score = 7

    # Acute admission flag (boolean)
    acute_score = 3 if acute_flag else 0

    # Charlson comorbidity score
    if charlson == 0:
        charlson_score = 0
    elif charlson == 1:
        charlson_score = 1
    elif charlson == 2:
        charlson_score = 2
    elif charlson == 3:
        charlson_score = 3
    else:
        charlson_score = 5

    # Emergency department visits in the last 6 months
    if ed_visits == 0:
        ed_score = 0
    elif ed_visits == 1:
        ed_score = 1
    elif ed_visits == 2:
        ed_score = 2
    elif ed_visits == 3:
        ed_score = 3
    else:
        ed_score = 4

    # Total LACE score
    score = los_score + acute_score + charlson_score + ed_score

    return score


```

<!-- grade_id: bmi-test, points: 10, grade: true, solution: false -->
```python
# Testing the function
assert LACE(4, False, 1, 0) == 5
assert LACE(4, True, 4, 7) == 16
```

```python

```

### 14.4 Care Management Criteria

Care managers use LACE as part of the criteria for assigning a care coodinator to a patient who has been recently discharged. If the score is above 10, then a care coordinator will be assigned. The other criteria they use is if the patient has been discharged with a diagnosis of CHF or COPD.  If the diagnosis field has CHF or COPD in it, then the patient will have a care coordinator assigned.

For this exercise, write another function that takes the same inputs as LACE() plus another diagnosis parameter, and return True or False depending on if the paient needs a care coordinator.

**NOTE** Pay attention to the fact that the order of parameters in this function definition are not the same as the order they were in the LACE score.  

<!-- grade_id: weight-answer, points: 0, grade: false, solution: true -->
```python
def assign_care_coordinator(diagnosis_cd, ed_visits, length_of_stay, acute_flag, charlson):
    """ (str, int, int, bool, int) -> bool
    Determines if a patient needs a care coordinator based on the LACE score
    and diagnosis of CHF or COPD.

    >>> assign_care_coordinator('None', 0, 4, False, 1)
    False

    >>> assign_care_coordinator('CHF', 0, 4, False, 1)
    True

    >>> assign_care_coordinator('COPD', 0, 4, False, 1)
    True

    >>> assign_care_coordinator('None', 7, 4, True, 4)
    True

    >>> assign_care_coordinator('CHF', 7, 4, True, 4)
    True
    """
    # Check if the diagnosis is CHF or COPD
    if diagnosis_cd in ['CHF', 'COPD']:
        return True

    # Calculate the LACE score
    lace_score = LACE(length_of_stay, acute_flag, charlson, ed_visits)

    # Assign a care coordinator if the LACE score is above 10
    if lace_score > 10:
        return True

    # Otherwise, no care coordinator is assigned
    return False




```

<!-- grade_id: weight-tests, points: 5, grade: true, solution: false -->
```python
# Testing the function
assert assign_care_coordinator('None', 0, 4, False, 1) == False
assert assign_care_coordinator('CHF', 0, 4, False, 1) == True
assert assign_care_coordinator('COPD', 0, 4, False, 1) == True
assert assign_care_coordinator('None', 7, 4, True, 4) == True
assert assign_care_coordinator('CHF', 7, 4, True, 4) == True
```

```python

```

### 14.5 qCSI COVID-19 Severity Index

See: https://www.mdcalc.com/quick-covid-19-severity-index-qcsi#evidence

Calculate the total risk score as per the point values assigned to respiratory rate, pulse oximetry, and O2 flow rate.  Then calculate and return the Risk Level.

In addition to the rules provided at the link above, also add the following checks for valid values:
* If `respiratory_rate <= 0` then return _invalid respiratory rate_
* If `pulse_ox <= 0` then return _invalid pulse ox_
* If `pulse_ox > 100` then return _invalid pulse ox_
* If `os_flow <= 0` then return _invalid O2 flow rate_

```python
def qcsi(respiratory_rate, pulse_ox, o2_flow):
    """(int, int, int) -> str
    * respiratory_rate is an integer value
    * pulse_ox is an integer value (e.g. 30 means 30%)
    * o2_flow is an integer value
    """

    if respiratory_rate <= 0:
        return "invalid respiratory rate"
    if pulse_ox <= 0 or pulse_ox > 100:
        return "invalid pulse ox"
    if o2_flow <= 0:
        return "invalid O2 flow rate"

    risk_score = 0

    # Assign points based on respiratory rate
    if respiratory_rate < 12:
        risk_score += 0
    elif 12 <= respiratory_rate <= 20:
        risk_score += 1
    elif 21 <= respiratory_rate <= 29:
        risk_score += 2
    else:  # respiratory_rate >= 30
        risk_score += 3

    # Assign points based on pulse oximetry
    if pulse_ox >= 95:
        risk_score += 0
    elif 90 <= pulse_ox < 95:
        risk_score += 1
    elif 80 <= pulse_ox < 90:
        risk_score += 2
    else:  # pulse_ox < 80
        risk_score += 3

    # Assign points based on O2 flow rate
    if o2_flow == 0:
        risk_score += 0
    elif o2_flow <= 2:
        risk_score += 1
    elif o2_flow <= 4:
        risk_score += 2
    else:  # o2_flow > 4
        risk_score += 3

    # Determine risk level based on total risk score
    if risk_score <= 2:
        risk_level = "Low Risk"
    elif 3 <= risk_score <= 5:
        risk_level = "Moderate Risk"
    else:  # risk_score >= 6
        risk_level = "High Risk"

    return risk_level
```

```python
# Assert statement's
assert (qcsi(29, 95, 1) == 'low-risk')
assert (qcsi(20, 93, 1) == 'low')
assert (qcsi(29, 88, 1) == 'high-intermediate')
assert (qcsi(29, 88, 4) == 'high')
assert (qcsi(30, 90, 1) == 'low-intermediate')
assert (qcsi(28, 92, -1) == 'invalid O2 flow rate')
assert (qcsi(22, 0, 4) == 'invalid pulse ox')
assert (qcsi(0, 97, 2) == 'invalid respiratory rate')
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-104-b6b631837b7c> in <cell line: 2>()
      1 # Assert statement's
----> 2 assert (qcsi(29, 95, 1) == 'low-risk')
      3 assert (qcsi(20, 93, 1) == 'low')
      4 assert (qcsi(29, 88, 1) == 'high-intermediate')
      5 assert (qcsi(29, 88, 4) == 'high')

AssertionError:
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week03/week03_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**



---